Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Discussion 0

This is the first (or 0th) discussion of Data 100. It covers basic Python, calculus, and probability concepts. 

## Question 1

For this question, write your answer on paper.

### Part A

Find the value of $x$ that minimizes the function $f(x) = x^2 + 4x + 4$.


<!-- grade_id: cell-d126d5df48b2218e, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

### Part B

Calculate the partial derivative of the following expression with respect to $x$.

$$f(x, y) = xy + \sin(x^2y) + \ln(x^3y)$$

**Note**: The partial derivative of a function with respect to a variable `x` involves taking the derivative of the function with respect to ``x`` while holding all other variables constant.


<!-- grade_id: cell-32ccafc40f19f143, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

### Part C

Now, calculate the partial derivative of the experssion with respect to $y$.

<!-- grade_id: cell-d28d0c3606c88e50, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 2

For this question, write your answer on paper.

**Are you smarter than a doctor? ONLY 46% OF DOCTORS GOT THIS QUESTION RIGHT.**

100 out of 10,000 women at age forty who participate in routine screening have breast cancer. 80 of those 100 women with breast cancer test positive. 950 out of 9,900 women without breast cancer also test positive. If 10,000 women in this age group undergo a routine screening, about what fraction of these women with positive tests will actually have breast cancer?

<!-- grade_id: cell-55b0f98a0de33933, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 3

Suppose we have the following list `lst` and array `arr`.

```python
import numpy as np
arr = np.arange(5)
lst = list(range(5))
```

```python
arr
```

```python
lst
```

What will be the output of the following lines of Python/Numpy code? Try to predict the output before running the cell.

```python
arr + 5
```

```python
lst + 5
```

```python
arr * 2
```

```python
lst * 2
```

## Question 4

***Note:*** The line `raise NotImplementedError()` indicates that the implementation still needs to be added. This is an exception derived from `RuntimeError`. Please comment out that line when you have implemented the function.


### Part A

Write a function that returns a list of numbers such that $ x_i=i^2 $ for $1\leq i \leq n$. Don't worry about the case where $n \leq 0$.

<!-- grade_id: cell-9388c587d0e639ff, points: 0, grade: false, solution: true -->
```python
def squares(n):
    """Compute the squares of numbers from 1 to n such that the ith element of the returned list equals i^2."""
    # YOUR CODE HERE
    raise NotImplementedError()
```

Your function should print `[1, 4, 9, 16, 25, 36, 49, 64, 81, 100]` for $n=10$.  
Check that it does:

```python
squares(10)
```

Check that squares returns the correct output for several inputs.

<!-- grade_id: squares-tests, points: 0, grade: true, solution: false -->
```python
assert squares(1) == [1]
assert squares(2) == [1, 4]
assert squares(10) == [1, 4, 9, 16, 25, 36, 49, 64, 81, 100]
assert squares(11) == [1, 4, 9, 16, 25, 36, 49, 64, 81, 100, 121]
```

### Part B

Evaluate the following summation (you may use code):

$$\sum_{i=1}^{100} i^3 + 3 i^2$$

<!-- grade_id: cell-e0447d4ce1ef9176, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
q2b_sum = ...
```

Check that your sum is correct.

```python
assert q2b_sum == 26517550
```

### Part C

Write a function `map_func` that will implement mapping and filtering on a list of values `list_vals`. This function will also take in a parameter `mapper`, which takes an input `x` and maps that value to a new value of the same type. The last parameter `filter_func` takes an input `y` and returns a boolean value based on if `y` satifies a certain condtion. In short, `map_func` should return a list of values from `n` that satisfy the condition established by `filter_func`.

**Note**: If you want to see examples of map_func used, look at the tests below.

<!-- grade_id: cell-c2f4ccfbf9d99081, points: 0, grade: false, solution: true -->
```python
def map_func(list_vals, mapper, filter_func):
    """Maps and filters the input list n based on the condition established by the input param filter_func. Return a list 
    containing elements of n that filter_func returns true on."""
    # YOUR CODE HERE
    raise NotImplementedError()
```

Check that `map_func` returns the correct output for several inputs.

```python
assert map_func([1, 2, 3], lambda x: x*2, lambda x: x > 1) == [4, 6]
assert map_func([], lambda x: x*1000, lambda x: x > -10) == []
assert map_func(["piglet", "gavin", "jim", "andy"], lambda x: x[:0:-1], lambda x: len(x) < 6) == ["niva", "mi", "ydn"]
```

## Question 5

### Part A

Write a function which takes in a string and returns `True` if the string is a palindrome. (A string is a palindrome if it is the same forwards and backwards.)

<!-- grade_id: cell-b1fa3ce3a6566fbc, points: 0, grade: false, solution: true -->
```python
def is_palindrome(word):
    """Return True if word is a palindrome."""
    # YOUR CODE HERE
    raise NotImplementedError()
```

Your function should return true for "racecar".

```python
is_palindrome("racecar")
```

Check that the function works for several inputs.

<!-- grade_id: test-is-palindrome, points: 0, grade: true, solution: false -->
```python
assert is_palindrome("aviddiva") == True
assert is_palindrome("clearlynotapalindrome") == False
assert is_palindrome("kayak") == True
assert is_palindrome("ab") == False
assert is_palindrome("abb") == False
assert is_palindrome("a") == True
```

### Part B

Write a function that flattens a nested Python list.

<!-- grade_id: cell-0335b7d85baac809, points: 0, grade: false, solution: true -->
```python
def flatten(lst):
    """Flattens the input list lst so that there are no nested lists."""
    # YOUR CODE HERE
    raise NotImplementedError()
```

Check that the function works for several inputs.

```python
assert flatten([1, 2, 3]) == [1, 2, 3]
assert flatten([1, 2, [3, 4]]) == [1, 2, 3, 4]
assert flatten([1, 2, [3, [4]], 5]) == [1, 2, 3, 4, 5]
assert flatten([1, 2, [3, [4]], [5, [[6]]]]) == [1, 2, 3, 4, 5, 6]
assert flatten([1, 2, [3, [[4]]], [5, [[6]]], [[[7]]]]) == [1, 2, 3, 4, 5, 6, 7]
```

Although there is no built-in `flatten` function for Python lists, numpy arrays have a `flatten` method that can be used to flatten an array. For example, to flatten a numpy array `arr` we would call `arr.flatten()`.


## Submission

You're done!

Before submitting this assignment, ensure to:

1. Restart the Kernel (in the menubar, select Kernel->Restart & Run All)
2. Validate the notebook by clicking the "Validate" button

Finally, make sure to **submit** the assignment via the Assignments tab in Datahub